<!DOCTYPE html>


<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    
    <title>502 Waiver Exam Page &mdash; Runestone Interactive Overview</title>
    
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="../_static/activecode.css" type="text/css" />
    <link rel="stylesheet" href="../_static/fitb.css" type="text/css" />
    <link rel="stylesheet" href="../_static/clickable.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pytutor.css" type="text/css" />
    <link rel="stylesheet" href="../_static/modal-basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/datafile.css" type="text/css" />
    <link rel="stylesheet" href="../_static/dragndrop.css" type="text/css" />
    <link rel="stylesheet" href="../_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="../_static/livecode.css" type="text/css" />
    <link rel="stylesheet" href="../_static/parsons.css" type="text/css" />
    <link rel="stylesheet" href="../_static/lib/prettify.css" type="text/css" />
    <link rel="stylesheet" href="../_static/poll.css" type="text/css" />
    <link rel="stylesheet" href="../_static/tabbedstuff.css" type="text/css" />
    <link rel="stylesheet" href="../_static/video.css" type="text/css" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    
    <script type="text/javascript">
      var DOCUMENTATION_OPTIONS = {
        URL_ROOT:    '../',
        VERSION:     '0.0',
        COLLAPSE_INDEX: false,
        FILE_SUFFIX: '.html',
        HAS_SOURCE:  true
      };
    </script>
    <script type="text/javascript" src="../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../_static/jquery.js"></script>
    <script type="text/javascript" src="../_static/underscore.js"></script>
    <script type="text/javascript" src="../_static/doctools.js"></script>
    <script type="text/javascript" src="../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../_static/codemirror.js"></script>
    <script type="text/javascript" src="../_static/xml.js"></script>
    <script type="text/javascript" src="../_static/css.js"></script>
    <script type="text/javascript" src="../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../_static/python.js"></script>
    <script type="text/javascript" src="../_static/javascript.js"></script>
    <script type="text/javascript" src="../_static/activecode.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/clike.js"></script>
    <script type="text/javascript" src="../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../_static/animationbase.js"></script>
    <script type="text/javascript" src="../_static/mchoice.js"></script>
    <script type="text/javascript" src="../_static/timedmc.js"></script>
    <script type="text/javascript" src="../_static/fitb.js"></script>
    <script type="text/javascript" src="../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../_static/timed.js"></script>
    <script type="text/javascript" src="../_static/clickable.js"></script>
    <script type="text/javascript" src="../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../_static/pytutor.js"></script>
    <script type="text/javascript" src="../_static/codelens.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/datafile.js"></script>
    <script type="text/javascript" src="../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../_static/livecode.js"></script>
    <script type="text/javascript" src="../_static/clike.js"></script>
    <script type="text/javascript" src="../_static/lib/jquery.min.js"></script>
    <script type="text/javascript" src="../_static/lib/jquery-ui.min.js"></script>
    <script type="text/javascript" src="../_static/lib/jquery.ui.touch-punch.min.js"></script>
    <script type="text/javascript" src="../_static/lib/prettify.js"></script>
    <script type="text/javascript" src="../_static/lib/underscore-min.js"></script>
    <script type="text/javascript" src="../_static/lib/lis.js"></script>
    <script type="text/javascript" src="../_static/parsons_setup.js"></script>
    <script type="text/javascript" src="../_static/parsons.js"></script>
    <script type="text/javascript" src="../_static/parsons-noconflict.js"></script>
    <script type="text/javascript" src="../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../_static/poll.js"></script>
    <script type="text/javascript" src="../_static/reveal.js"></script>
    <script type="text/javascript" src="../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="https://cdn.mathjax.org/mathjax/latest/MathJax.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="top" title="Runestone Interactive Overview" href="../index.html" />
    <link rel="next" title="Practice Assessment for 502 Waiver" href="PracticeWaiver502.html" />
    <link rel="prev" title="Dictionary operations" href="../Assessment/Dictionaryoperations.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'BasicsofPython';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = false;
  eBookConfig.basecourse = 'BasicsofPython';
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>



  </head>
  <body role="document">


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../index.html">Basics of Python</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-share-alt" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/BasicsofPython/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/BasicsofPython/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/static/help/bugreports.html'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">502 Waiver Exam Page</a><ul>
<li><a class="reference internal" href="#waiver-exam">502 Waiver Exam</a></li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - Dictionary operations" data-toggle="tooltip">
    <a href="../Assessment/Dictionaryoperations.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Practice Assessment for 502 Waiver' data-toggle="tooltip" >
    <a href="PracticeWaiver502.html" >
      <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="waiver-exam-page">
<h1>502 Waiver Exam Page<a class="headerlink" href="#waiver-exam-page" title="Permalink to this headline">¶</a></h1>
<p>The following is the assessment for the 502 Waiver. This should take about 30 minutes to complete. Content based on the first few chapters will include:</p>
<ul class="simple">
<li>Assignment Statements</li>
<li>Conditional Statements</li>
<li>Indexing</li>
<li>For Loops</li>
<li>Dictionaries</li>
</ul>
<div class="section" id="waiver-exam">
<h2>502 Waiver Exam<a class="headerlink" href="#waiver-exam" title="Permalink to this headline">¶</a></h2>
<p><strong>Instructions:</strong> For questions that ask you to write code, write the code you want to save in the provided active code windows. The last code you have saved for each problem by the deadline is what will be graded. Instructions are comments in the code window. <em>Do not</em> hit <strong>Finish Exam</strong> until you are finished.</p>

    <ul data-component="timedAssessment" data-time=30 id="timed1"   >
    
<textarea data-component="activecode" id=waiver502_1 data-lang="python"    data-timelimit=25000   data-audio=''     >
# Create a variable called uni and assign it the value of
# this phrase: University of Michigan



=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(uni, "University of Michigan", "Either uni has not been named correctly or has not been assigned the value of the string, 'University of Michigan'")

myTests().main()
</textarea>

    <pre data-component="parsons" id="waiver502_2" data-maxdist=5>
        <span data-question>waiver502_1: Arrange the following code snipits in the correct order to loop through the list and print out only those that are integers, otherwise print None.
</span>full_list = ["Hello", 5, "world", -7.3,
            ["umich", "edu"], -1817, 0.000,
             "runestone"]
---
for item in full_list:
---
   if type(item) == int:
---
      print item
---
   else:
---
      print "None"</pre>
    
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="waiver502_3">
            <p>waiver502_2: In the following for loop, what does <code class="docutils literal"><span class="pre">item</span></code> represent?</p>
<div class="highlight-python"><div class="highlight"><pre><span class="k">for</span> <span class="n">item</span> <span class="ow">in</span> <span class="nb">range</span><span class="p">(</span><span class="mi">3</span><span class="p">,</span><span class="mi">15</span><span class="p">):</span>
</pre></div>
</div>

            <li data-component="answer"  id="waiver502_3_opt_a">Items that you need for creating the list produced by the range function.</li><li data-component="feedback" id="waiver502_3_opt_a">The only input that the range function needs is between the parenthesis in the function call, item is not used for that.</li>
            
            <li data-component="answer" data-correct id="waiver502_3_opt_b">The elements of the list produced by the range function call.</li><li data-component="feedback" id="waiver502_3_opt_b">Elem represents each element in the list that is created by the range function call.</li>
            
            <li data-component="answer"  id="waiver502_3_opt_c">Nothing, that is an improper use of a for loop.</li><li data-component="feedback" id="waiver502_3_opt_c">The range function can be used as the object that is iterated over in a for loop.</li>
            

            </ul>
            
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="waiver502_4">
            waiver502_3: What is type of the value of the (human) fifth position in this list: [&#8220;School&#8221;, &#8220;Ann Arbor&#8221;, &#8220;93.0&#8221;, 38, [6], &#8220;Information&#8221;, {3: &#8220;three&#8221;, &#8220;Six&#8221;: 6, 109: 109.0}]?
            <li data-component="answer"  id="waiver502_4_opt_a">Integer</li><li data-component="feedback" id="waiver502_4_opt_a">The human fifth position in the list refers to [6], which is a list. The contents of the list are a subsection and their type is not accounted for when refering to the entire list object.</li>
            
            <li data-component="answer"  id="waiver502_4_opt_b">Float</li><li data-component="feedback" id="waiver502_4_opt_b">The human fifth position in the list refers to [6], which is a list.</li>
            
            <li data-component="answer"  id="waiver502_4_opt_c">String</li><li data-component="feedback" id="waiver502_4_opt_c">The human fifth position in the list refers to [6], which is a list.</li>
            
            <li data-component="answer" data-correct id="waiver502_4_opt_d">List</li><li data-component="feedback" id="waiver502_4_opt_d">The human fifth position in the list refers to [6], which is a list.</li>
            
            <li data-component="answer"  id="waiver502_4_opt_e">Dictionary</li><li data-component="feedback" id="waiver502_4_opt_e">The human fifth position in the list refers to [6], which is a list, as denoted by the square [] brackets.</li>
            

            </ul>
            
    <pre data-component="parsons" id="waiver502_5" data-maxdist=5>
        <span data-question>waiver502_4: Arrange the following code snipits so that var is printed out 8 times, each time on a new line.
</span>var = "9"
---
for pos in range(8):
---
   print var</pre>
    
<textarea data-component="activecode" id=waiver502_6 data-lang="python"    data-timelimit=25000   data-audio=''     >
# If the type of the last variable in lst is a float or an integer,
# then assign the value of the_type to be True, otherwise assign
# it to be False. the_type is currently assigned the value None.

lst = [["Dog", "cat", 8], "pet", {"Lassie": "dog", "Porky": "pig"}, 10, 9, 8.99]
the_type = None

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(the_type, True, "The if/else conditional statement has not been set up correctly or the value of the_type is not a Boolean.")

myTests().main()
</textarea>

<textarea data-component="activecode" id=waiver502_7 data-lang="python"    data-timelimit=25000   data-audio=''     >
# Write code to assign the variable abc the value that is produced
# when the length of a string is the input of the range function.
# In this case, when the length of the string "Hello" is used,
# a list is produced and is saved as the value of the variable abc.

wrd = "Hello"


=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(abc, [0, 1, 2, 3, 4], "abc has not been assigned the correct value. There may be impropper use of the len or range functions.")

myTests().main()
</textarea>

<textarea data-component="activecode" id=waiver502_8 data-lang="python"    data-timelimit=25000   data-audio=''     >
# Write a statement to assign the second to last word in this
# string to the variable called scri.

sent = "The Michigan Difference is a wonderful motto."

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(scri, "wonderful", "scri has not been correctly assigned the value of 'wonderful', possibly due to incorrect indexing.")

myTests().main()
</textarea>

            <ul data-component="multiplechoice" data-multipleanswers="false"  id="waiver502_9">
            waiver502_5: After the following Python code is executed, what will the the type of L[4:5] if L = [&#8216;p&#8217;, &#8216;y&#8217;, &#8216;t&#8217;, &#8216;h&#8217;, &#8216;o&#8217;, &#8216;n&#8217;]?
            <li data-component="answer"  id="waiver502_9_opt_a">String</li><li data-component="feedback" id="waiver502_9_opt_a">Due to the way that L was indexed in to, the resulting type will be a list.</li>
            
            <li data-component="answer" data-correct id="waiver502_9_opt_b">List</li><li data-component="feedback" id="waiver502_9_opt_b">Due to the way that L was indexed in to, the resulting type will be a list.</li>
            
            <li data-component="answer"  id="waiver502_9_opt_c">Dictionary</li><li data-component="feedback" id="waiver502_9_opt_c">Unless the element that is indexed in to is a dictionary, the type of the value will not be a dictionary. In this case, even though the element is a string, the way that it was indexed will result in a list with the single element.</li>
            

            </ul>
            
<textarea data-component="activecode" id=waiver502_10 data-lang="python"    data-timelimit=25000   data-audio=''     >
# Write code so that if the letter q is in the string
# then the program should assign the string "q is present"
# to a variable called sent_ans, otherwise, it should assign
# that variable the value of "q is not present".

sent_long = "Programming is a useful skill for people to know in this day and age, where computers are an important and vital tool in societies."

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(sent_ans, "q is not present", "Testing whether sent_ans has been correctly set up and assigned the right value.")

myTests().main()
</textarea>

            <ul data-component="multiplechoice" data-multipleanswers="false"  id="waiver502_11">
            <p>waiver502_6: What will print out at the end of this section of code?</p>
<div class="highlight-python"><div class="highlight"><pre><span class="n">first</span> <span class="o">=</span> <span class="mi">3</span>
<span class="n">sec</span> <span class="o">=</span> <span class="mi">2</span>
<span class="n">first</span> <span class="o">==</span> <span class="n">sec</span> <span class="o">+</span> <span class="mi">1</span>
<span class="k">print</span> <span class="n">first</span>
</pre></div>
</div>

            <li data-component="answer"  id="waiver502_11_opt_a">6</li><li data-component="feedback" id="waiver502_11_opt_a">First was compared against sec, not combined with sec + 1, it is still 3.</li>
            
            <li data-component="answer" data-correct id="waiver502_11_opt_b">3</li><li data-component="feedback" id="waiver502_11_opt_b">First still is assigned the value of 3, even after the comparison against sec + 1.</li>
            
            <li data-component="answer"  id="waiver502_11_opt_c">False</li><li data-component="feedback" id="waiver502_11_opt_c">First is not reassigned the value of the comparison of first and sec + 1, which would have been True, not False, so first is still 3.</li>
            
            <li data-component="answer"  id="waiver502_11_opt_d">True</li><li data-component="feedback" id="waiver502_11_opt_d">First is not reassigned the value of the comparison of first and sec + 1, so it cannot be True. First is still 3.</li>
            
            <li data-component="answer"  id="waiver502_11_opt_e">None of the above, there will be an error.</li><li data-component="feedback" id="waiver502_11_opt_e">No error occurs, all of the lines are permissible in python.</li>
            

            </ul>
            
    <pre data-component="parsons" id="waiver502_12" data-maxdist=5>
        <span data-question>waiver502_7: Arrange the following code snipits in the correct order so that the list of letters creates a dictionary where in the key is the letter and the value is the number of times it has occured in the list.
</span>let_list = ['a', 't',
      'm', 'i','c', 'i',
      'b', 'e', 'y', 'p',
      't', 's', 'c', 'a',
      'l', 'i', 'a', 'y',
      's', 's', 'u', 'e',
      'e', 'd',]
dict = {}
---
for let in let_list:
---
   if let in dict:
---
      dict[let] += 1
---
   else:
---
      dict[let] = 1</pre>
    
<textarea data-component="activecode" id=waiver502_13 data-lang="python"    data-timelimit=25000   data-audio=''     >
# Using the dictionary provided, add an extra
# key value pair, whose key is the string festival
# and whose value is the string summer fest.

word_dict = {"yellow": "maize", "blue": "same thing", "NQ": "North Quad", "UMSI": "University of Michigan School of Information", "A2": "Ann Arbor", "Big House": "The Football Stadium", "Language": "Python", "Fishbowl": "Computing Site"}

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(sorted(word_dict), sorted({"yellow": "maize", "blue": "same thing", "NQ": "North Quad", "UMSI": "University of Michigan School of Information", "A2": "Ann Arbor", "Big House": "The Football Stadium", "Language": "Python", "Fishbowl": "Computing Site", "festival": "summer fest"}), "Checking whether word_dict has the right keys.")
      self.assertEqual(sorted(word_dict.values()), sorted("maize", "same thing", North Quad", "University of Michigan School of Information", "Ann Arbor", "The Football Stadium", "Python", "Computing Site", "summer fest"), "Checking whether word_dict has the right values.")


myTests().main()
</textarea>

            <ul data-component="multiplechoice" data-multipleanswers="false"  id="waiver502_14">
            <p>waiver502_8: What would occur when this code is printed?</p>
<div class="highlight-python"><div class="highlight"><pre><span class="n">d</span> <span class="o">=</span> <span class="p">{}</span>
<span class="n">d</span><span class="p">[</span><span class="s">&#39;a&#39;</span><span class="p">]</span> <span class="o">=</span> <span class="mi">5</span>
<span class="n">d</span><span class="p">[</span><span class="s">&#39;b&#39;</span><span class="p">]</span> <span class="o">=</span> <span class="mi">4</span>
<span class="n">d</span><span class="p">[</span><span class="mi">3</span><span class="p">]</span> <span class="o">=</span> <span class="s">&#39;b&#39;</span>
<span class="k">print</span> <span class="mi">4</span> <span class="ow">in</span> <span class="n">d</span>
</pre></div>
</div>

            <li data-component="answer"  id="waiver502_14_opt_a">True</li><li data-component="feedback" id="waiver502_14_opt_a">Using the in opperator on dictionaries checks the keys, so since 4 is not a key, it will return False.</li>
            
            <li data-component="answer" data-correct id="waiver502_14_opt_b">False</li><li data-component="feedback" id="waiver502_14_opt_b">Using the in opperator on dictionaries checks the keys, so since 4 is not a key, it will return False.</li>
            
            <li data-component="answer"  id="waiver502_14_opt_c">'b'</li><li data-component="feedback" id="waiver502_14_opt_c">When using the in opperator on dictionaries, it checks the keys and would not find the 4. Additionally, the in opperator returns a boolean value, not the key associated with a value.</li>
            
            <li data-component="answer"  id="waiver502_14_opt_d">4</li><li data-component="feedback" id="waiver502_14_opt_d">When using the in opperator on dictionaries, it checks the keys and would not find the 4. Additionally, the in opperator returns a boolean value, not the value it checks for.</li>
            
            <li data-component="answer"  id="waiver502_14_opt_e">Error occurs</li><li data-component="feedback" id="waiver502_14_opt_e">The in opperator is permissible in python, and returns a boolean value.</li>
            

            </ul>
            
    <pre data-component="parsons" id="waiver502_15" data-maxdist=5>
        <span data-question>waiver502_9: Arrange the code snipits so that the values of the dictionary are added together
so that the total number is displayed in a string that reads "This is my 29th
favorite day of the month." and is printed out.
</span>num_dict = {"University of Michigan": 19, "month": 4,
           "student": 3.5, "season": 2, "question": .5}

num = 0
---
for key in num_dict:
---
   num = num + num_dict[key]
---
print "This is my " + num + "th favorite day of the month."</pre>
    
            <ul data-component="multiplechoice" data-multipleanswers="true"  id="waiver502_16">
            <p>waiver502_10: What is wrong with this dictionary?</p>
<div class="highlight-python"><div class="highlight"><pre><span class="p">{</span><span class="s">&quot;diction&quot;</span><span class="p">:</span> <span class="n">ary</span><span class="p">,</span> <span class="s">&quot;gaming consoles&quot;</span><span class="p">:</span> <span class="p">[</span><span class="s">&quot;ps4&quot;</span><span class="p">,</span> <span class="s">&quot;gamecube&quot;</span><span class="p">,</span> <span class="s">&quot;xbox360&quot;</span><span class="p">,</span> <span class="s">&quot;nintendo64&quot;</span><span class="p">,</span> <span class="s">&quot;Atari&quot;</span><span class="p">],</span> <span class="s">&quot;phones&quot;</span><span class="p">:</span> <span class="mi">3</span> <span class="s">&quot;broken&quot;</span><span class="p">:</span> <span class="mi">10</span><span class="p">,</span> <span class="n">donut</span><span class="p">:</span> <span class="p">{</span><span class="s">&quot;Glazed&quot;</span><span class="p">:</span> <span class="mi">10</span><span class="p">,</span> <span class="mi">2</span><span class="p">:</span> <span class="s">&quot;blueberry&quot;</span><span class="p">,</span> <span class="s">&quot;Chocolate&quot;</span><span class="p">:</span> <span class="s">&quot;Many&quot;</span><span class="p">}}</span>
</pre></div>
</div>

            <li data-component="answer"  id="waiver502_16_opt_a">There is a list and a dictionary inside of the dictionary</li><li data-component="feedback" id="waiver502_16_opt_a">Not quite right, you are allowed to have dictionaries and lists inside of other lists and dictionaries.</li>
            
            <li data-component="answer" data-correct id="waiver502_16_opt_b">There is no assignment statement</li><li data-component="feedback" id="waiver502_16_opt_b">Yes, there is no assignment statement, thus there is no way to use the dictionary later on.</li>
            
            <li data-component="answer" data-correct id="waiver502_16_opt_c">Some of the values and keys are missing quotes</li><li data-component="feedback" id="waiver502_16_opt_c">Yes, some of the values and keys that are supposed to be strings are missing quotes around them, thus the dictionary will think that they are missing variables and cause an error.</li>
            
            <li data-component="answer" data-correct id="waiver502_16_opt_d">There's a comma missing</li><li data-component="feedback" id="waiver502_16_opt_d">Yes, there is a comma missing, which would cause an error when trying to run. Syntax is important to running programs.</li>
            
            <li data-component="answer"  id="waiver502_16_opt_e">a number is a key, which is not allowed</li><li data-component="feedback" id="waiver502_16_opt_e">Not quite right, numbers are allowed to be keys and values.</li>
            

            </ul>
            
<textarea data-component="activecode" id=waiver502_17 data-lang="python"    data-timelimit=25000   data-audio=''     >
# assign to the variable desti, strings that describe
# whether or not the length of "Pip install" is above 15,
# below 8, or inbetween. If it is above 15, assign it
# "The lenth of the string is greater than 15",
# if it is below 8, assign it "The lenth of the string is
# less than 8", and if it is between, assign it
# "The length of the string is between 15 and 8".

check = "Pip install"

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(desti, "The length of the string is between 15 and 8", "desti may have been assigned the incorrect string.")

myTests().main()
</textarea>

    <pre data-component="parsons" id="waiver502_18" data-maxdist=5>
        <span data-question>waiver502_11: Arrange the code snipits so that the list of integers is iterated through so that if the resulting value is divisable by 3 with zero remaining, then "div by 3: " and the number is printed, otherwise nothing happens.
</span>lst_nums = [9, 392, 4, 10, 244, 206, 506, 1983, 1817, 16]
---
for nums in lst_nums:
---
   if num % 3 == 0:
---
      print "div by 3: ", num</pre>
    
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="waiver502_19">
            waiver502_12: Can you index into a dictionary?
            <li data-component="answer"  id="waiver502_19_opt_a">Yes, you can index into a dictionary by using their positions.</li><li data-component="feedback" id="waiver502_19_opt_a">Dictionaries don't have positions that you can use to index.</li>
            
            <li data-component="answer"  id="waiver502_19_opt_b">Yes, you can index into a dictionary because they are like lists, just with {} brackets instead of [] ones.</li><li data-component="feedback" id="waiver502_19_opt_b">Dictionaries aren't quite like lists, they are unordered whereas lists are ordered.</li>
            
            <li data-component="answer" data-correct id="waiver502_19_opt_c">No, you cannot index into a dictionary because they do not have positions.</li><li data-component="feedback" id="waiver502_19_opt_c">You cannot index into a dictionary because thery lack positions to use for numeric indices.</li>
            
            <li data-component="answer"  id="waiver502_19_opt_d">No, you cannot index into a dictionary because the {} brackets are different than then [] ones.</li><li data-component="feedback" id="waiver502_19_opt_d">They are different kinds of brackets, but that's not the reason why you can't index into a dictionary.</li>
            

            </ul>
            
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="waiver502_20">
            <p>waiver502_13: What is printed by the following statement?</p>
<div class="highlight-python"><div class="highlight"><pre><span class="n">new_dict</span> <span class="o">=</span> <span class="p">{</span><span class="s">&quot;Pizza&quot;</span><span class="p">:</span> <span class="mi">7</span><span class="p">,</span> <span class="s">&quot;Large&quot;</span><span class="p">:</span> <span class="mi">2</span><span class="p">,</span> <span class="s">&quot;Medium&quot;</span><span class="p">:</span> <span class="s">&quot;Zero&quot;</span><span class="p">,</span> <span class="s">&quot;Extra-large&quot;</span><span class="p">:</span> <span class="s">&quot;5&quot;</span><span class="p">}</span>
<span class="k">print</span> <span class="n">new_dict</span><span class="p">[</span><span class="s">&quot;Medium&quot;</span><span class="p">]</span>
</pre></div>
</div>

            <li data-component="answer"  id="waiver502_20_opt_a">"Zero"</li><li data-component="feedback" id="waiver502_20_opt_a">Not quite, the quotes will not be printed out.</li>
            
            <li data-component="answer" data-correct id="waiver502_20_opt_b">Zero</li><li data-component="feedback" id="waiver502_20_opt_b">The value of the key will be printed out without quotes.</li>
            
            <li data-component="answer"  id="waiver502_20_opt_c">Error, you cannot index into a dictionary.</li><li data-component="feedback" id="waiver502_20_opt_c">No, this is not indexing into a dictionary, this is refering to the value that is associated with the key.</li>
            

            </ul>
            
<textarea data-component="activecode" id=waiver502_21 data-lang="python"    data-timelimit=25000   data-audio=''     >
# Write code so that the length of the list
# is assigned to a variable called student_len.

students = [52, 45, 90, 18, 150, 390, 1029]


=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(student_len, len(students), "You incorrectly assigned the value of student_len to be the length of students.")

myTests().main()ac
</textarea>

<textarea data-component="activecode" id=waiver502_22 data-lang="python"    data-timelimit=25000   data-audio=''     >
# Write code to change the number of students
# who are in Information so that there are 490 students.

enrollment = {"Architecture": 638, "InterArts": 15, "Law": 885, "Pharmacy": 420, "Information": 459, "LSA": 18290, "Engineering": 8723, "Social Work": "472"}
</textarea>

            <ul data-component="multiplechoice" data-multipleanswers="true"  id="waiver502_23">
            <p>waiver502_14: Which of these are equivalent? (Choose all that apply)</p>
<div class="highlight-python"><div class="highlight"><pre><span class="n">var</span> <span class="o">=</span> <span class="s">&quot;Go Blue!&quot;</span>
</pre></div>
</div>

            <li data-component="answer" data-correct id="waiver502_23_opt_a">var[4:5]</li><li data-component="feedback" id="waiver502_23_opt_a">This refers to the substring "l", not "lu", because the slice goes up to but does include position 5.</li>
            
            <li data-component="answer" data-correct id="waiver502_23_opt_b">var[4]</li><li data-component="feedback" id="waiver502_23_opt_b">This refers to the substring "l".</li>
            
            <li data-component="answer"  id="waiver502_23_opt_c">var["l"]</li><li data-component="feedback" id="waiver502_23_opt_c">An Error occurs because when indexing, you must use the numerical position.</li>
            
            <li data-component="answer" data-correct id="waiver502_23_opt_d">var[-3]</li><li data-component="feedback" id="waiver502_23_opt_d">This refers to the substring "g".</li>
            

            </ul>
            
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="waiver502_24">
            <p>waiver502_15: What would be saved to the value of detention at the end of the code?</p>
<div class="highlight-python"><div class="highlight"><pre><span class="n">tardy</span> <span class="o">=</span> <span class="mi">4</span>
<span class="n">late</span> <span class="o">=</span> <span class="mi">8</span>
<span class="n">detention</span> <span class="o">=</span> <span class="s">&quot;&quot;</span>

<span class="k">if</span> <span class="n">late</span> <span class="o">&gt;</span> <span class="n">tardy</span><span class="p">:</span>
   <span class="n">detention</span> <span class="o">=</span> <span class="s">&quot;You have been late more than tardy.&quot;</span>
<span class="k">elif</span> <span class="n">late</span> <span class="o">==</span> <span class="n">tardy</span><span class="p">:</span>
   <span class="n">dentention</span> <span class="o">=</span> <span class="s">&quot;You are just as late as you are tardy.&quot;</span>
<span class="k">else</span><span class="p">:</span>
   <span class="n">dentention</span> <span class="o">=</span> <span class="s">&quot;You are more tardy than late.&quot;</span>

<span class="k">if</span> <span class="n">tardy</span> <span class="o">&gt;=</span> <span class="n">late</span><span class="p">:</span>
   <span class="n">detention</span> <span class="o">=</span> <span class="s">&quot;You might be more tardy.&quot;</span>
<span class="k">else</span><span class="p">:</span>
   <span class="n">detention</span> <span class="o">=</span> <span class="s">&quot;You&#39;re less tardy.&quot;</span>
</pre></div>
</div>

            <li data-component="answer"  id="waiver502_24_opt_a">"You have been late more than tardy."</li><li data-component="feedback" id="waiver502_24_opt_a">While that did happen at first, the program continues and this is overwritten.</li>
            
            <li data-component="answer" data-correct id="waiver502_24_opt_b">"You're less tardy."</li><li data-component="feedback" id="waiver502_24_opt_b">Correct!</li>
            
            <li data-component="answer"  id="waiver502_24_opt_c">"You are more tardy than late."</li><li data-component="feedback" id="waiver502_24_opt_c">Not exactly, not only does late happen to be greater than tardy, but there is another if else conditional that the program goes through.</li>
            

            </ul>
            
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="waiver502_25">
            waiver502_16: How would you use the range function to create a list that spans from 23 up to and including 89?
            <li data-component="answer" data-correct id="waiver502_25_opt_a">range(23, 90)</li><li data-component="feedback" id="waiver502_25_opt_a">Range, when given two inputs, will be inclusive of the first and exclusive of the second, so this will work!</li>
            
            <li data-component="answer"  id="waiver502_25_opt_b">range(22, 90)</li><li data-component="feedback" id="waiver502_25_opt_b">Range, when given two inputs, will be inclusive of the first and exclusive of the second.</li>
            
            <li data-component="answer"  id="waiver502_25_opt_c">range(23, 89)</li><li data-component="feedback" id="waiver502_25_opt_c">Range, when given two inputs, will be inclusive of the first and exclusive of the second.</li>
            
            <li data-component="answer"  id="waiver502_25_opt_d">range(22, 89)</li><li data-component="feedback" id="waiver502_25_opt_d">Range, when given two inputs, will be inclusive of the first and exclusive of the second.</li>
            

            </ul>
            </ul>
    </div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - Dictionary operations' data-toggle="tooltip">
    <a href="../Assessment/Dictionaryoperations.html" >
      <i class='prevNav glyphicon glyphicon-chevron-left'></i>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="../Assessment/Dictionaryoperations.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Practice Assessment for 502 Waiver' data-toggle="tooltip" >
    <a href="PracticeWaiver502.html" >
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right'></i>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="PracticeWaiver502.html" >Next Section - Practice Assessment for 502 Waiver</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2015 Lauren.
      Created using <a href="http://sphinx.pocoo.org/">Sphinx</a> 1.3.1.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>